```python
from autograder_term2week8 import *
_globals = globals()
import math
from copy import copy
import pickle
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki, Matt Woolway</center>

# <center> Worksheet 12: Sets, Frozensets, Dictionaries, Files</center>

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-da7dc1d6c8be001a, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

Output:
```
You've changed the contents of the file; thank you!
Your username has been recognised; thank you!

5 out of 5 marks
5
```

## Question 1

### (i)

(a) Create a list, a tuple, a set and a frozenset, each containing the odd integers between 13 and 103 inclusive.

<!-- grade_id: cell-5bad39258340d53f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list1 = list(range(13,104,2))
tuple1 = tuple(range(13,104,2))
set1 = set(range(13,104,2))
fs1 = frozenset(range(13,104,2))
### END SOLUTION

print(list1)
print(tuple1)
print(set1)
print(fs1)
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103]
(13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103)
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103}
frozenset({13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103})
```

<!-- grade_id: cell-985d45de372b4792, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ia(_globals)
```

Output:
```
list1 contains the correct data
tuple1 contains the correct data
set1 contains the correct data
fs1 contains the correct data

list1 is a list
tuple1 is a tuple
set1 is a set
fs1 is a frozenset

2 out of 2
2
```

(b) Set up `list2`, identically equal to `list1`, and `list3`, a copy of `list1` created using the `copy` function.
Also set up `set2`, identically equal to `set1`, and `set3`, a copy of `set1` created using the `copy` function.

<!-- grade_id: cell-2c5a79c99b279f52, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list2 = list1
list3 = copy(list1)
set2 = set1
set3 = copy(set1)
### END SOLUTION

print(list2)
print(list3)
print(set2)
print(set3)
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103]
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103]
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103}
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103}
```

<!-- grade_id: cell-e2dc6aa11630389b, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ib(_globals)
```

Output:
```
list2 contains the same elements as list1
list2 is the same object as list1, which is correct

list3 contains the same elements as list1
list3 is not the same object as list1, which is correct

set2 contains the same elements as set2
set2 is the same object as set1, which is correct

set3 contains the same elements as set1
set3 is not the same object as set1, which is correct

4 out of 4
4
```

(c) Introduce the element 105 to `list1` using the `append` method, and introduce the element 105 to `set1` using the `add` method. 

<!-- grade_id: cell-db56a67357ca11dd, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list1.append(105)
set1.add(105)
### END SOLUTION

print(list1)
print(set1)
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105]
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105}
```

<!-- grade_id: cell-6dcd5a8637a9e567, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ic(_globals)
```

Output:
```
list1 is correct
set1 is correct
2 out of 2
2
```

(d) Check the values of `list2`, `list3`, `set2` and `set3`. Which of the following is true?

<ol>
    <li>The values of <code>list2</code>, <code>list3</code>, <code>set2</code> and <code>set3</code> have all changed.</li>
    <li>The values of <code>list2</code>, <code>list3</code>, <code>set2</code> and <code>set3</code> have all stayed the same.</li>
    <li>The values of <code>list2</code> and <code>set2</code> have changed, but the values of <code>list3</code> and <code>set3</code> have stayed the same.</li>
    <li>The values of <code>list3</code> and <code>set3</code> have changed, but the values of <code>list2</code> and <code>set2</code> have stayed the same.</li>
    <li>Only<code>list2</code> has changed; the rest have have stayed the same.</li>
    <li>Only<code>set2</code> has changed; the rest have have stayed the same.</li>
</ol>

<!-- grade_id: cell-6db2e1ccdd6aed11, points: 0, grade: false, solution: true -->
```python
print(list2)
print(list3)
print(set2)
print(set3)

### BEGIN SOLUTION
question1id_answer = 3
### END SOLUTION
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105]
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103]
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105}
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103}
```

<!-- grade_id: cell-79f6b308a520727c, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1id_answer == question1id())
print('test case passed!')
```

Output:
```
test case passed!
```

(e) Again, introduce the element 105 to `list1` using the `append` method, and introduce the element 105 to `set1` using the `add` method.

<!-- grade_id: cell-53dbb31f38cee238, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list1.append(105)
set1.add(105)
### END SOLUTION

print(list1)
print(set1)
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105, 105]
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105}
```

<!-- grade_id: cell-2b3879d9477c0181, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ie(_globals)
```

Output:
```
list1 is correct
set1 is correct
2 out of 2
2
```

(f) Once more, check the values of `list2`, `list3`, `set2` and `set3`. Which of the following is true?

<ol>
    <li>The values of <code>list2</code>, <code>list3</code>, <code>set2</code> and <code>set3</code> have all changed.</li>
    <li>The values of <code>list2</code>, <code>list3</code>, <code>set2</code> and <code>set3</code> have all stayed the same.</li>
    <li>The values of <code>list2</code> and <code>set2</code> have changed, but the values of <code>list3</code> and <code>set3</code> have stayed the same.</li>
    <li>The values of <code>list3</code> and <code>set3</code> have changed, but the values of <code>list2</code> and <code>set2</code> have stayed the same.</li>
    <li>Only<code>list2</code> has changed; the rest have have stayed the same.</li>
    <li>Only<code>set2</code> has changed; the rest have have stayed the same.</li>
</ol>

<!-- grade_id: cell-27214580e10d517e, points: 0, grade: false, solution: true -->
```python
print(list2)
print(list3)
print(set2)
print(set3)

### BEGIN SOLUTION
question1if_answer = 5
### END SOLUTION
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105, 105]
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103]
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105}
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103}
```

<!-- grade_id: cell-bcc127bfac5cdadc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1if_answer == question1if())
print('test case passed!')
```

Output:
```
test case passed!
```

(g) Redefine `list1` as the list of multiples of 5 between 1005 and 1305 inclusive, and redefine `set1` as the set of multiples of 5 between 1005 and 1305 inclusive.

<!-- grade_id: cell-b4e82724652a1d66, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list1 = list(range(1005,1310,5))
set1 = set(range(1005,1310,5))
### END SOLUTION

print(list1)
print(set1)
```

Output:
```
[1005, 1010, 1015, 1020, 1025, 1030, 1035, 1040, 1045, 1050, 1055, 1060, 1065, 1070, 1075, 1080, 1085, 1090, 1095, 1100, 1105, 1110, 1115, 1120, 1125, 1130, 1135, 1140, 1145, 1150, 1155, 1160, 1165, 1170, 1175, 1180, 1185, 1190, 1195, 1200, 1205, 1210, 1215, 1220, 1225, 1230, 1235, 1240, 1245, 1250, 1255, 1260, 1265, 1270, 1275, 1280, 1285, 1290, 1295, 1300, 1305]
{1280, 1025, 1155, 1285, 1030, 1160, 1290, 1035, 1165, 1295, 1040, 1170, 1300, 1045, 1175, 1305, 1050, 1180, 1055, 1185, 1060, 1190, 1065, 1195, 1070, 1200, 1075, 1205, 1080, 1210, 1085, 1215, 1090, 1220, 1095, 1225, 1100, 1230, 1105, 1235, 1110, 1240, 1115, 1245, 1120, 1250, 1125, 1255, 1130, 1260, 1005, 1135, 1265, 1010, 1140, 1270, 1015, 1145, 1275, 1020, 1150}
```

<!-- grade_id: cell-352b4f20ea015272, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ig(_globals)
```

Output:
```
list1 is correct
set1 is correct
2 out of 2
2
```

(h) For a third time, check the values of `list2`, `list3`, `set2` and `set3`. Which of the following is true?

<ol>
    <li>The values of <code>list2</code>, <code>list3</code>, <code>set2</code> and <code>set3</code> have all changed.</li>
    <li>The values of <code>list2</code>, <code>list3</code>, <code>set2</code> and <code>set3</code> have all stayed the same.</li>
    <li>The values of <code>list2</code> and <code>set2</code> have changed, but the values of <code>list3</code> and <code>set3</code> have stayed the same.</li>
    <li>The values of <code>list3</code> and <code>set3</code> have changed, but the values of <code>list2</code> and <code>set2</code> have stayed the same.</li>
    <li>Only<code>list2</code> has changed; the rest have have stayed the same.</li>
    <li>Only<code>set2</code> has changed; the rest have have stayed the same.</li>
</ol>

<!-- grade_id: cell-f483f5b86d2c05d2, points: 0, grade: false, solution: true -->
```python
print(list2)
print(list3)
print(set2)
print(set3)

### BEGIN SOLUTION
question1ih_answer = 2
### END SOLUTION
```

Output:
```
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105, 105]
[13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103]
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103, 105}
{13, 15, 17, 19, 21, 23, 25, 27, 29, 31, 33, 35, 37, 39, 41, 43, 45, 47, 49, 51, 53, 55, 57, 59, 61, 63, 65, 67, 69, 71, 73, 75, 77, 79, 81, 83, 85, 87, 89, 91, 93, 95, 97, 99, 101, 103}
```

<!-- grade_id: cell-881cc6b1f89ed9af, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1ih_answer == question1ih())
print('test case passed!')
```

Output:
```
test case passed!
```

### (ii)

(a) Define `row1` as `[1, 3, 5]`, `row2` as `[1, 4, 7]` and `row3` as `[1, 5, 9]`, and define `list_of_lists1` as `[row1, row2, row3]`.

Set up `list_of_lists2` as identical to `list_of_lists1`, and `list_of_lists3` as a copy of it.

<!-- grade_id: cell-3385361e6bb5e74d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
row1 = [1, 3, 5]
row2 = [1, 4, 7]
row3 = [1, 5, 9]

list_of_lists1 = [row1, row2, row3]
list_of_lists2 = list_of_lists1
list_of_lists3 = copy(list_of_lists1)
### END SOLUTION

print(list_of_lists1)
print(list_of_lists2)
print(list_of_lists3)
```

Output:
```
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
```

<!-- grade_id: cell-5ed214164dad5d2a, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1iia(_globals)
```

Output:
```
list_of_lists1 contains the right values

list_of_lists2 contains the right values
list_of_lists2 is the same object as list_of_lists1, which is correct

list_of_lists3 contains the right values
list_of_lists3 is not the same object as list_of_lists1, which is correct

3 out of 3
3
```

(b) Type 
```python
list_of_lists1[1][1] = 400
```

Check the values of `list_of_lists2` and `list_of_lists3`. Which of the following is true?

<ol>
    <li>The values of <code>list_of_lists2</code> and <code>list_of_lists3</code> have both changed.</li>
    <li>The values of <code>list_of_lists2</code> and <code>list_of_lists3</code> have both stayed the same.</li>
    <li>The value of <code>list_of_lists2</code> has changed, but that of <code>list_of_lists3</code> has stayed the same.</li>
    <li>The value of <code>list_of_lists3</code> has changed, but that of <code>list_of_lists2</code> has stayed the same.</li>
</ol>

<!-- grade_id: cell-24d681b71db986bf, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list_of_lists1[1][1] = 400

print(list_of_lists1)
print(list_of_lists2)
print(list_of_lists3)

question1iib_answer = 1
### END SOLUTION
```

Output:
```
[[1, 3, 5], [1, 400, 7], [1, 5, 9]]
[[1, 3, 5], [1, 400, 7], [1, 5, 9]]
[[1, 3, 5], [1, 400, 7], [1, 5, 9]]
```

<!-- grade_id: cell-cf69c20d7081cea7, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iib_answer == question1iib())
print('test case passed!')
```

Output:
```
test case passed!
```

(c) Redefine `row1`, `row2`, `row3` and `list_of_lists1` as in (a) above. 

Create an object, `list_of_lists4`, that is a <b>deep</b> copy of `list_of_lists1`: that is, not only is it itself a copy, but all three of the elements that make it up are copies of the corresponding elements of `list_of_lists1`. It should remain the same in the face of changes to any of the numbers that form part of `list_of_lists1`.

<b>Note</b>: don't do this just by repeating the definition of `list_of_lists1`; find a smarter way! There are at least two ways of doing this task. One uses a comprehension, and the other uses another function from the `copy` module, whose name you might be able to guess (and which is in any case available on the Internet...)

<!-- grade_id: cell-49d7b1d42a394359, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
row1 = [1, 3, 5]
row2 = [1, 4, 7]
row3 = [1, 5, 9]

list_of_lists1 = [row1, row2, row3]
from copy import deepcopy
list_of_lists4 = deepcopy(list_of_lists1)
### END SOLUTION

print(list_of_lists1)
print(list_of_lists4)
```

Output:
```
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
```

<!-- grade_id: cell-9d44c0503ab37154, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1iic(_globals)
```

Output:
```
list_of_lists1 contains the right values

list_of_lists4 contains the right values

list_of_lists4 is not the same object as list_of_lists1, which is correct

list_of_lists4 is a deep copy: its rows are separate objects from those of list_of_lists1

4 out of 4
4
```

(d) Can you find the other method, calling your second deep copy `list_of_lists5?` (Autograding won't be able to tell if you've "cheated" here by doing the same thing twice; but you'll know!)

<!-- grade_id: cell-8797c7cf90bbc928, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
row1 = [1, 3, 5]
row2 = [1, 4, 7]
row3 = [1, 5, 9]

list_of_lists1 = [row1, row2, row3]
list_of_lists5 = [copy(row) for row in list_of_lists1]
### END SOLUTION

print(list_of_lists1)
print(list_of_lists5)
```

Output:
```
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
```

<!-- grade_id: cell-4e03943b224f0325, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1iid(_globals)
```

Output:
```
list_of_lists1 contains the right values

list_of_lists5 contains the right values

list_of_lists5 is not the same object as list_of_lists1, which is correct

list_of_lists5 is a deep copy: its rows are separate objects from those of list_of_lists1

4 out of 4
4
```

(e) Type 
```python
list_of_lists1[1][1] = 400
```

Check the value of `list_of_lists4`. Which of the following is true?

<ol>
    <li>The value of <code>list_of_lists4</code> has changed.</li>
    <li>The value of <code>list_of_lists4</code> has stayed the same.</li>
</ol>

<!-- grade_id: cell-dab59ec93525e1ad, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
list_of_lists1[1][1] = 400

print(list_of_lists1)
print(list_of_lists4)

question1iie_answer = 2
### END SOLUTION
```

Output:
```
[[1, 3, 5], [1, 400, 7], [1, 5, 9]]
[[1, 3, 5], [1, 4, 7], [1, 5, 9]]
```

<!-- grade_id: cell-e2827ff5f8f3b898, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1iie_answer == question1iie())
print('test case passed!')
```

Output:
```
test case passed!
```

## Question 2

### (i)

The following data consists of:
<ul>
    <li>Eight dictionaries, each with the same three keys: <code>'radius'</code>,  <code>'mass'</code> and <code>'escape_velocity'</code>; the value associated with each key is a float.</li>
    <li>One "master dictionary", consisting of eight keys corresponding to the names of the planets in our solar system; the value associated with each key is a dictionary.</li>
</ul>

```python
mercury_data = {'radius' : 2.4e6, 'mass' : 3.3e23, 
    'escape_velocity' : 4.3e3}
venus_data = {'radius' : 6.1e6, 'mass' : 4.9e24, 
    'escape_velocity' : 1.0e4}
earth_data = {'radius' : 6.3e6, 'mass' : 6.0e24, 
    'escape_velocity' : 1.1e4}
mars_data = {'radius' : 3.4e6, 'mass' : 6.4e23, 
    'escape_velocity' : 5.0e3}
jupiter_data = {'radius' : 6.9e7, 'mass' : 1.9e27, 
    'escape_velocity' : 6.0e4}
saturn_data = {'radius' : 5.7e7, 'mass' : 5.6e26, 
    'escape_velocity' : 3.6e4}
uranus_data = {'radius' : 2.5e7, 'mass' : 8.7e25, 
    'escape_velocity' : 2.2e4}
neptune_data = {'radius' : 2.5e7, 'mass' : 1.0e26, 
    'escape_velocity' : 2.4e4}
planet_data = {
    'mercury': mercury_data,
    'venus': venus_data,
    'earth': earth_data,
    'mars': mars_data,
    'jupiter': jupiter_data,
    'saturn': saturn_data,
    'uranus': uranus_data,
    'neptune': neptune_data,
}
```

(a) Using a comprehension, create a <b>list</b> `radius_list` containing all the planetary radii.

<!-- grade_id: cell-feef8e594c7c7108, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
radius_list = [x['radius'] for x in planet_data.values()]
### END SOLUTION

print(radius_list)
```

Output:
```
[2400000.0, 6100000.0, 6300000.0, 3400000.0, 69000000.0, 57000000.0, 25000000.0, 25000000.0]
```

<!-- grade_id: cell-b930ad8d3f909ab9, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ia(_globals)
```

Output:
```
radius_list is correct
Test passed!
1
```

(b) Using a comprehension, create a <b>set</b> `mass_set` containing all the planetary masses.

<!-- grade_id: cell-d57b551cbb142ba2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mass_set = {x['mass'] for x in planet_data.values()}
### END SOLUTION

print(mass_set)
```

Output:
```
{6e+24, 5.6e+26, 3.3e+23, 4.9e+24, 1.9e+27, 6.4e+23, 1e+26, 8.7e+25}
```

<!-- grade_id: cell-ab71aab9394ec18c, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ib(_globals)
```

Output:
```
mass_set is correct
Test passed!
1
```

(c) Create a <b>set</b> `name_set` containing all the names of the planets; this can be done without using a comprehension.

<!-- grade_id: cell-121cbc171d4dbdbe, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
name_set = set(planet_data)
### END SOLUTION

print(name_set)
```

Output:
```
{'uranus', 'earth', 'saturn', 'venus', 'mercury', 'mars', 'neptune', 'jupiter'}
```

<!-- grade_id: cell-487b52adfcefaaf2, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ic(_globals)
```

Output:
```
name_set is correct
Test passed!
1
```

(d) Using a comprehension, create a <b>dictionary</b> `ev_dict`, the keys of which are the names of the planets and the associated values of which are the escape velocities.

<!-- grade_id: cell-c7f53923a0d38513, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
ev_dict = {planet:data['escape_velocity'] 
           for planet, data in planet_data.items()}
### END SOLUTION

print(ev_dict)
```

Output:
```
{'mercury': 4300.0, 'venus': 10000.0, 'earth': 11000.0, 'mars': 5000.0, 'jupiter': 60000.0, 'saturn': 36000.0, 'uranus': 22000.0, 'neptune': 24000.0}
```

<!-- grade_id: cell-0ceb196e018d19cc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2id(_globals)
```

Output:
```
ev_dict is correct
Test passed!
1
```

### (ii)

The following is an attempt to set up a little dictionary linking the IUPAC systematic and preferred designations of the isomers of xylene (dimethylbenzene). 

```python
xylene_isomers = {
    [1, 2] : 'orthoxylene', 
    [1, 3] : 'metaxylene',
    [1, 4] : 'paraxylene'}
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[64], line 1
----> 1 xylene_isomers = {
      2     [1, 2] : 'orthoxylene', 
      3     [1, 3] : 'metaxylene',
      4     [1, 4] : 'paraxylene'}

TypeError: unhashable type: 'list'
```

The way it's meant to work is that since, for example, 1,2-dimethylbenzene is also known as orthoxylene, I can type
```python
xylene_isomers([1, 2])
```
and get the output `'orthoxylene'`.

(a) However, it doesn't work. Fix it! Write a new version of `xylene_isomers` in which the keys are data structures (don't use strings or ints for example), but ones that work.

<!-- grade_id: cell-45514ee84187b5cb, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
xylene_isomers = {
    (1, 2) : 'orthoxylene', 
    (1, 3) : 'metaxylene',
    (1, 4) : 'paraxylene'}
### END SOLUTION

print(xylene_isomers)
```

<!-- grade_id: cell-bfe12f7e0d2af965, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2iia(_globals)
```

(b) Which of the following data types and data structures can serve as the keys in a dictionary? Type your answer as a set.

<ol>
    <li>floats;</li>
    <li>ints;</li>
    <li>complexes;</li>
    <li>strings;</li>
    <li>lists;</li>
    <li>tuples;</li>
    <li>NumPy arrays;</li>
    <li>sets;</li>
    <li>frozensets;</li>
    <li>dictionaries.</li>
</ol>

<!-- grade_id: cell-dda9e8f6cee7e834, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
question2iib_answer = {1,2,3,4,6,9}
### END SOLUTION
```

<!-- grade_id: cell-d38702d62a5b38d6, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2iib_answer == question2iib())
print('test case passed!')
```

(c) Which of the following data types and data structures can serve as the values in a dictionary? Type your answer as a set.

<ol>
    <li>floats;</li>
    <li>ints;</li>
    <li>complexes;</li>
    <li>strings;</li>
    <li>lists;</li>
    <li>tuples;</li>
    <li>NumPy arrays;</li>
    <li>sets;</li>
    <li>frozensets;</li>
    <li>dictionaries.</li>
</ol>

<!-- grade_id: cell-2d4bd5bb56172f4b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
question2iic_answer = {1,2,3,4,5,6,7,8,9,10}
### END SOLUTION

question2iic_answer = # insert a set of ints
```

<!-- grade_id: cell-164cdda160462d87, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2iic_answer == question2iic())
print('test case passed!')
```

## Question 3

Write and test a function called `power_set` which takes as its argument a set, and returns the set of that set's subsets (each subset being represented as a frozenset). So, for example, 
```python
power_set({0,1})
```
should return `{frozenset(), frozenset({1}), frozenset({0}), frozenset({0, 1})}`.

(I think this is a naturally <b>recursive</b> task. I also think this is quite tricky.)

Note that the empty set is represented in Python by `set()` rather than by `{}`, which represents the empty dictionary.

<!-- grade_id: cell-7fc4f6ccc3b3f268, points: 0, grade: false, solution: true -->
```python
def power_set(this_set):
    ### BEGIN SOLUTION
    """
    Returns the power set of this_set, the various subsets being
    represented as frozensets
    """
    # base case
    if this_set == set():
        return {frozenset()}
    # recursion st
    else:        
        # remove an element
        el = this_set.pop()
        
        # power set of that set
        short_power_set = power_set(this_set)
        
        # copy with the last element appended to each
        short_power_set2 = {subset | frozenset({el}) for subset in short_power_set}
        
        # union of the two is returned
        return short_power_set | short_power_set2
    ### END SOLUTION
```

<!-- grade_id: cell-18689f9543ab4e90, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3(_globals)
```

## Question 4
(a) Take your `xylene_isomers` dictionary, as corrected, and convert it into a string, `xylene_isomers_string`.

<!-- grade_id: cell-de31117e983b3356, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
xylene_isomers_string = str(xylene_isomers)
### END SOLUTION

print(xylene_isomers_string)
```

<!-- grade_id: cell-a2b91a558fe53dbb, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4a(_globals)
```

(b) Write this string to an external file called `xylene_isomers.txt`. Autograding can't help you here!

<!-- grade_id: cell-0ed1707326f6838b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
fo = open('xylene_isomers.txt','w')
fo.write(xylene_isomers_string)
fo.close()
### END SOLUTION
```

(c) Read in your external file as a string, `xylene_isomers_string2`.

<!-- grade_id: cell-23cccfc73b453766, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
fo = open('xylene_isomers.txt', 'r')
xylene_isomers_string2 = fo.read()
fo.close
### END SOLUTION

print(xylene_isomers_string2)
```

<!-- grade_id: cell-d7d9639e634f1bbc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4c(_globals)
```

(d) Convert this string back into a dictionary, `xylene_isomers2`.

<!-- grade_id: cell-a35bcdc8e815203b, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
xylene_isomers2 = eval(xylene_isomers_string2)
### END SOLUTION
print(xylene_isomers2)
```

<!-- grade_id: cell-2646ba379df0ad0b, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4d(_globals)
```

(e) Take your `xylene_isomers` dictionary, as corrected, and pickle it as `xylene_isomers.pickle`. Autograding can't help here!

<!-- grade_id: cell-44ce9d96abd7ed04, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
po = open('xylene_isomers.pickle', 'wb')
pickle.dump(xylene_isomers, po)
po.close()
### END SOLUTION
```

(f) Unpickle this file as `xylene_isomers3`.

<!-- grade_id: cell-dde568c125cd5488, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
po = open('xylene_isomers.pickle', 'rb')
xylene_isomers3 = pickle.load(po)
po.close()
### END SOLUTION

xylene_isomers3
```

<!-- grade_id: cell-72571ecba9ebed63, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4f(_globals)
```

(g) Pickle the `cos` function from `math`, and unpickle it as `new_cos`. Test it.

<!-- grade_id: cell-b67b7fd7dda67d98, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
po = open('cos.pickle', 'wb')
pickle.dump(math.cos, po)
po.close()

po = open('cos.pickle', 'rb')
new_cos = pickle.load(po)
po.close()
### END SOLUTION

print(new_cos(math.pi))
```

<!-- grade_id: cell-581fb821b12fb1c7, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4g(_globals)
```

(h) Pickle your `power_set` function, and unpickle it as `new_power_set`. Test it.

<!-- grade_id: cell-b05283bfcdf07066, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
po = open('power_set.pickle', 'wb')
pickle.dump(power_set, po)
po.close()

po = open('power_set.pickle', 'rb')
new_power_set = pickle.load(po)
po.close()
### END SOLUTION

print(new_power_set({0, 1, 2}))
```

<!-- grade_id: cell-e2b127a78684912f, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4h(_globals)
```

```python

```