```python
from autograder_term2week5 import *
_globals = globals()
import sympy as sp
sp.init_printing()
import numpy as np
import math
import sympy.plotting as splt
import matplotlib.pyplot as plt
%matplotlib inline
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki, Matt Woolway</center>

# <center> Worksheet 10: SymPy</center>

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-edd6c6fe992ac654, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

## Question 1(i)

(a) Define `x` and `t` as SymPy symbolic variables.

<!-- grade_id: cell-3fc307878e1a9e93, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

(x, t)
```

<!-- grade_id: cell-46337e404d685821, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ia(_globals)
```

(b) Define `expr` as $(x-2)(x+1)(x-(2+i))(x-((2-i)).$ Use SymPy complex numbers.

<!-- grade_id: cell-bf426dcc39f3a4ee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

expr
```

<!-- grade_id: cell-da6e34ce755c07d0, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ib(_globals)
```

(c) Expand `expr` to give `expr2`.

<!-- grade_id: cell-483a179309921041, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

expr2
```

<!-- grade_id: cell-c438c178b78dc472, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ic(_globals)
```

(d) Find the zeros of `expr2` using `solve`; call the resulting object `expr2_zeros`.

<!-- grade_id: cell-275ab15d204f7c11, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

expr2_zeros
```

<!-- grade_id: cell-5007b7fef608950b, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1id(_globals)
```

(e) Find, as an exact SymPy rational, the integral of `expr2` between $-3$ and $3$; call this `int_expr`.

<!-- grade_id: cell-e6b86a8e3a14e459, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

int_expr
```

<!-- grade_id: cell-7572bb4ef99c4950, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ie(_globals)
```

(f) Set up `f` as an undefined SymPy function of $t$, and create an expression called `fexpr` which is $f(t)$ substituted into `expr2`.

<!-- grade_id: cell-377f72b45c5c86d4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

(f, fexpr)
```

<!-- grade_id: cell-81ba14e2362cb0d9, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1if(_globals)
```

(g) Given that $x=f(t)$, find the general solution `gen_sol` of the ordinary differential equation

$$\frac{dx}{dt} = x^4-5x^3+7 x^2+3x-10,$$

as an expression in $f(t)$. (Don't panic if this takes quite a long time to execute.)

<!-- grade_id: cell-5793e81f8df0d50e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

gen_sol
```

<!-- grade_id: cell-da84d476dc222ea4, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1ig(_globals)
```

## Question 1 (ii)

Autograding won't help for any of this!

Using the functions `plot_parametric`, `plot_implicit`, `plot3d_parametric_line` and `plot3d_parametric_surface`, create plots of:

(a) The Lissajous curve $x=\cos(5\theta), y= \sin(3\theta)$ for $0\le \theta \le 2\,\pi$.

(b) The implicit curve $y^2 = x^3-x$ for $-3\le x \le 3$.

(c) A helix in 3d (just a line, not a surface).

(d) The surface of a torus (ring doughnut or bagel shape).

(e) A helix winding round a torus (just show the line, not the surface).

<!-- grade_id: cell-fe82240418492676, points: 0, grade: false, solution: true -->
```python
# cell for part (a)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d6b67b71f213baa7, points: 0, grade: false, solution: true -->
```python
# cell for part (b)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-246db9ce9a41786f, points: 0, grade: false, solution: true -->
```python
# cell for part (c)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8cbcc36f51eb8fa2, points: 0, grade: false, solution: true -->
```python
# cell for part (d)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-07818c747e91e690, points: 0, grade: false, solution: true -->
```python
# cell for part (e)
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 2 (i)

(a) Define variables $p=6$ and $q=9$ <em>as SymPy integers</em>, and calculate <em>using SymPy functions and rationals</em> the quantity
$$x=\left(\frac{q}{2}+\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}+\left(\frac{q}{2}-\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}.$$
Verify Cardano's Theorem, that $x$ is a solution of the cubic
$$x^3=p\,x+q.$$

<!-- grade_id: cell-fe73392b99dcc38d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

(x, x**3, p*x+q)
```

<!-- grade_id: cell-c4a58b62e61622f0, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ia(_globals)
```

(b) Repeat for $p=3$ and $q=18$.

<!-- grade_id: cell-f81eea0fb1edb157, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

(x, x**3, p*x+q)
```

<!-- grade_id: cell-5727229a7d29a2da, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ib(_globals)
```

(c) Which of the following is/are true?

<ol>
    <li>The values of $x^3$ and $p\,x+q$ are not equal; this is a counterexample to the theorem.</li>
    <li>The values of $x^3$ and $p\,x+q$ are equal, as is revealed by typing <code>x ** 3 == p*x + q</code>.</li>
    <li>The values of $x^3$ and $p\,x+q$ are equal, as is revealed by typing <code>sp.simplify(x ** 3) == sp.simplify(p*x + q)</code>.</li>
    <li>The values of $x^3$ and $p\,x+q$ are equal, as is revealed by typing <code>sp.simplify(x ** 3 - (p*x + q)) == 0</code>.</li>
    <li>The values of $x^3$ and $p\,x+q$ are equal, as is revealed by typing <code>(x ** 3).equals(p*x + q)</code>.</li>
</ol>

Type in the correct answers enclosed in curly brackets: for example, if you think statements 1, 3 and 4 are true, type `{1, 3, 4}`.

<!-- grade_id: cell-e3c878f032c8325e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

q2ic_answer = # give the correct options in curly brackets
```

<!-- grade_id: cell-1839617dea822581, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ic(_globals)
```

(d) The only real root of $x^3=3x+18$ is $x=3$. Which of the following is true?

<ol>
    <li>The numerical value of <code>x</code> found in this way is not $3$.</li>
    <li>The numerical value of <code>x</code> found in this way is $3$, but this appears in an unhelpful form; however, if we simplify it using <code>simplify</code>, it becomes 3.</li>
    <li>The numerical value of <code>x</code> found in this way is $3$, but this appears in an unhelpful form; moreover, if we simplify it using <code>simplify</code>, it does not become 3.</li>
</ol>

<!-- grade_id: cell-3e52f326323eb208, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

q2id_answer = # insert one of 1, 2 or 3
```

<!-- grade_id: cell-ec841bfa442874f5, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q2id_answer == question2id())
print('test case passed!')
```

The `simplify` subsubmodule of the `simplify` submodule of SymPy has a specialised function called `nthroot`, which attempts to compute a real $n$th root of a sum of surds; you can test it by typing, for example,

```python
from sympy.simplify.simplify import nthroot
nthroot(-25+22*sp.sqrt(2), 3)
```

which should deliver a simple expression for the real cube root of $-25+22\,\sqrt{2}$.

(e) Use this function instead of raising to fractional powers in order to calculate, for $p=3$ and $q=18$, the exact value of
$$x=\left(\frac{q}{2}+\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}+\left(\frac{q}{2}-\sqrt{\left[\frac{q}{2}\right]^2-\left[\frac{p}{3}\right]^3}\right)^{1/3}.$$

<!-- grade_id: cell-b93cdde30dc5adaa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

(x, x**3, p*x+q)
```

<!-- grade_id: cell-0864f75752234de0, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2ie(_globals)
```

(f) Investigate the case $p=6$, $q=4$, for which the cubic has three real roots. Which of the following is true?

<ol>
    <li>Cardano's method, used in this way with rational powers, finds none of the real roots.</li>
    <li>Cardano's method, used in this way with rational powers, finds just one of the real roots, in a nicely simplified form.</li>
    <li>Cardano's method, used in this way with rational powers, finds just one of the real roots, in an unhelpful form; it can be found in a nicely simplified form by using <code>nthroot</code> instead.</li>
    <li>Cardano's method, used in this way with either rational powers or <code>nthroot</code>, finds just one of the real roots, in an unhelpful form; it can be found in a nicely simplified form by using <code>expand</code> with the keyword argument <code>complex</code> set to <code>True</code>.</li>
</ol>

<!-- grade_id: cell-80798a6234dd0146, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

(x,(x**3).expand(), p*x+q)
```

<!-- grade_id: cell-ba67312ebdc82ed6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

q2if_answer = # insert one of 1, 2, 3, 4
```

<!-- grade_id: cell-2e05b4aa4b216e78, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q2if_answer == question2if())
print('test case passed!')
```

## Question 2 (ii)

Every real number $x$ can be expressed as a unique <b>continued fraction</b>,

$$x = a_1 + \frac{1}{a_2+\frac{1}{a_3+\frac{1}{a_4+\dots}}}$$

We can represent this by the sequence

$$a_1, a_2, a_3, a_4, \dots$$

If $x$ is rational then this sequence is finite; otherwise it is infinite.

The following function is for calculating continued fractions of positive real numbers, truncating if necessary after the sequence attains length `n`:

```python
def continued_fraction(x, n):
    """
    Calculates the continued fraction of x up to a maximum of n terms
    """
    cf = []
    
    for i in range(n):
        intx = int(x)
        
        cf.append(intx)
        
        if x == intx:
            return cf
        
        x = 1/(x-intx)
        
    return cf
```

(a) Calculate continued fractions for the SymPy values of $\pi$, $\sqrt{2}$, $\phi=(1+\sqrt{5})/2$ and $22/7$ (as a Rational) with `n` set to 10. Call these, respectively, `pi_cf`, `sqrt2_cf`, `phi_cf` and `rat_cf`.

<!-- grade_id: cell-1c236919c7c27450, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(pi_cf)
print(sqrt2_cf)
print(phi_cf)
print(rat_cf)
```

<!-- grade_id: cell-1b35f93bb54360b5, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2iia(_globals)
```

(b) Calculate continued fractions for the SymPy, NumPy and `math` values of $\pi$, with `n` set to 50. Which of the following is true?

<ol>
    <li>All three values of $\pi$ give the same continued fraction.</li>
    <li>One of the three values of $\pi$ gives a different continued fraction from the others, and this is the most reliable.</li>
    <li>One of the three values of $\pi$ gives a different continued fraction from the others, and this is the least reliable.</li>
</ol>

<!-- grade_id: cell-0b9c39fec595a7f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

q2iib_answer = # insert one of 1, 2 or 3
```

<!-- grade_id: cell-a9212ce720886bce, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(q2iib_answer == question2iib())
print('test case passed!')
```

(c) Write and test a function called `from_continued_fraction`, which takes as its single argument a list of positive ints `[a1, a2, a3,  ..., an]` and returns, as a SymPy rational, the value of 

$$x = a_1 + \frac{1}{a_2+\frac{1}{a_3+\frac{1}{a_4+\dots}}}$$

up to $a_n$.

<!-- grade_id: cell-416fbeb88a921ee4, points: 0, grade: false, solution: true -->
```python
def from_continued_fraction(cf):
    # YOUR CODE HERE
    raise NotImplementedError()
float(from_continued_fraction([3,7,15]))
```

<!-- grade_id: cell-28fb6cdd3c5afffb, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2iic(_globals)
```

(d) Write and test a function called `rational_sequence`, which takes as its arguments a number $x$ and a positive int $n$, and returns a list of successive rational approximations to $x$ obtained by successive truncations of $x$'s continued fraction. 

This list should be the same length as `continued_fraction(x, n)`.

As a test

```python
rational_sequence(sp.sqrt(2), 5)
```

should return `[1, 3/2, 7/5, 17/12, 41/29]`.

<!-- grade_id: cell-5c5f3f1391268a52, points: 0, grade: false, solution: true -->
```python
def rational_sequence(x, n):
    # YOUR CODE HERE
    raise NotImplementedError()
print(rational_sequence(sp.pi,5))
```

<!-- grade_id: cell-1078a527f92575d0, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2iid(_globals)
```

(e) Plot, on the same pair of axes, the convergence of the rational sequences for the SymPy values of $\pi$, $\sqrt{2}$ and $\phi = (1+\sqrt{5})/2$, each with $n=10$. In each case the horizontal coordinates should be $0, 1, 2, \dots, 9.$

<!-- grade_id: cell-cfeea56da0c4037f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b0a75c427b72d32, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2iie_plot(_globals)
```

## Question 3

(a) Set up `x` as a SymPy symbol, with the keyword argument `real` set to `True`, so SymPy knows `x` must be real.

<!-- grade_id: cell-d26d2a7d86dbf401, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ddd3903fba829355, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

(b) Define `y` as the symbolic polynomial expression $x^4-8 x^3+26 x^2-40 x+21$.

<!-- grade_id: cell-48fc6f85ba0fdd3d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

y
```

<!-- grade_id: cell-47e77a00fdfd379f, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3b(_globals)
```

(c) Use the `plot` function from `sympy.plotting` to plot `y` for values of `x` between $-1$ and $5$. Autograding won't help here!

<!-- grade_id: cell-1ca2aac4b9b0effc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

(d) Using `solve`, find the zeros of `y` as SymPy quantities, as `y_zeros`. If you've set `x` and `y` up correctly, these should all be real.

<!-- grade_id: cell-1c2b3dd4b6821228, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

y_zeros
```

<!-- grade_id: cell-117dc8f5c6abb617, points: 0, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3d(_globals)
```

(e) Using `lambdify`, create and test a function called `yfunc`, compatible with NumPy arrays, that maps `x` to $x^4-8 x^3+26 x^2-40 x+21$.

<!-- grade_id: cell-9d7ea7bb4d5afca6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

yfunc
```

<!-- grade_id: cell-dc4e4daa4fa507d4, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3e(_globals)
```

(f) Hence, create a plot of `y` against `x`, for 101 equally-spaced values of `x` between $-1$ and $5$, that uses only the `plot` function from `matplotlib.pyplot`.

<!-- grade_id: cell-4896b9b9d19945a8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-16deb245b20ad8fc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3f_plot(_globals)
```

(g) Create a plot of `y`, for the same values of `x` as above, showing the real zeros as big red dots. Include lines for the horizontal and vertical axes. 

<!-- grade_id: cell-6d7e0c37b16ef8d1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1e096121be9787d5, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3g_plot(_globals)
```

(h) Using `diff` and `solve`, find, as a Python list, the $x$-coordinates `x_stat` of the (real) stationary points of `y`. Also find, as a Python list, the associated $y$-coordinates, `y_stat`.

<!-- grade_id: cell-0da9684fd0e8fe24, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print((x_stat, y_stat))
```

<!-- grade_id: cell-6f2019e2136270a5, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3h(_globals)
```

(j) Create an ordinary `matplotlib` plot of `y` showing the real stationary points as big red dots. Include lines for the horizontal and vertical axes. Again, use 101 equally-spaced values of `x` between $-1$ and $5$.

<!-- grade_id: cell-440f910b7170e401, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7277b1a9170266e9, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3j_plot(_globals)
```

(k) Write and test a function called `plot_with_sps` that takes as its arguments a symbolic expression `y` representing a function of `x`, and a tuple `xrange` containing a symbolic variable `x` (which may be assumed to be real) and maximum and minimum values for plotting. 

Your function should return no output, but should instead create a graphic in which a plot of `y` against `x` is shown, together with the real stationary points as big red dots. For example

```python
x = sp.symbols('x', real=True)
plot_with_sps(x**4-8*x**3+26*x**2-40*x+21, (x, -1, 5))
```

should generate the graphic from part (j).

It's tricky to get autograding to work for all the possible ways you might do this!

<!-- grade_id: cell-42de3b080fa8deb7, points: 0, grade: false, solution: true -->
```python
def plot_with_sps(y, xrange):
    # YOUR CODE HERE
    raise NotImplementedError()
x = sp.symbols('x', real=True)
plot_with_sps(x**4-8*x**3+26*x**2-40*x+21, (x, -1, 5))
```

(l) (<b>Challenging, maybe? Though it can be done by tweaking one line!</b>): Adapt your `plot_with_sps` code so that it works even if `x` has not been set up as an explicitly <em>real</em> symbol; so that for example

```python
x = sp.symbols('x')
plot_with_sps(x**4-8*x**3+26*x**2-40*x+21, (x, -1, 5))
```

should still generate the graphic from part (j). You'll need to think how to get rid of non-real roots.

<!-- grade_id: cell-4daab71a9b80aecc, points: 0, grade: false, solution: true -->
```python
def plot_with_sps(y, xrange):
    # YOUR CODE HERE
    raise NotImplementedError()
x = sp.symbols('x')
plot_with_sps(x**4-8*x**3+26*x**2-40*x+21, (x, -1, 5))
```

## Question 4

(a) Write a function called `newton_sympy` which takes as its arguments

<ul>
    <li>a <b>symbolic expression</b> <code>expr</code>;</li>
    <li>a <b>symbolic variable</b> <code>var</code>;</li>
    <li>a starting value, <code>x0</code>;</li>
    <li>a number of iterations, <code>niterate</code>;</li>
    <li>an keyword-only argument, <code>method</code>, whose default value is the string <code>'symbolic'</code>.</li>
</ul>

It should then:

<ul>
    <li>differentiate <code>expr</code> symbolically with respect to <code>var</code>, and hence derive, in <code>var</code>, the function <code>newt_expr</code> that must be iterated in order to solve <code>expr == 0</code> using Newton's method;</li>
    <li>iterate the function specified by <code>newt_expr</code> a total of <code>niterate</code> times, returning the final iterate.</li>
</ul>

Exactly what happens during the iteration will need to be considered quite carefully. 

<ul>
    <li>If <code>method</code> is set to <code>'symbolic'</code>, its default value, then each new iterate should simply be obtained by substituting the previous one into <code>newt_expr</code> using <code>subs</code>;</li>
    <li>If <code>method</code> is set to <code>'simplify'</code>, then each new iterate should be obtained by substituting the previous one into <code>newt_expr</code> using <code>subs</code>, and then simplifying the resulting expression;</li>
    <li>If <code>method</code> is set to <code>'float'</code>, then each new iterate should be obtained by substituting the previous one into <code>newt_expr</code> using <code>subs</code>, and then casting the resulting expression as a core Python float;</li>
    <li>If <code>method</code> is set to <code>'complex'</code>, then each new iterate should be obtained by substituting the previous one into <code>newt_expr</code> using <code>subs</code>, and then casting the resulting expression as a core Python complex;</li>
    <li>If <code>method</code> is set to <code>'lambdify'</code>, then <code>subs</code> should not be used at all; instead, a function <code>newt_expr_func</code> should be created using <code>lambdify</code>, and each iterate obtained by evaluating this function on the previous iterate.</li>
</ul>

<!-- grade_id: cell-ba7b053b47667219, points: 0, grade: false, solution: true -->
```python
def newton_sympy(expr, var, x0, niterate, *, method='symbolic'):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-585d34f45ccf6d50, points: 6, grade: true, solution: false -->
```python
# 6 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4a(_globals)
```

(b) Use your `newton_sympy` function, with 5 iterations starting at the int 1 and the default `method` setting, to find a good rational approximation `sqrt2_frac` to $\sqrt{2}$. [Hint: what is a simple equation satisfied by $\sqrt{2}$?]

<!-- grade_id: cell-f8dfa184fc5316dd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(sqrt2_frac)
print((float(sqrt2_frac),np.sqrt(2)))
```

<!-- grade_id: cell-15dc9f36d29e3131, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4b(_globals)
```

(c) Use your `newton_sympy` function with `method` set to `float` or `complex` as appropriate, to find the roots of $z^3-1=0$ associated with the starting values $2.0$, $-1+i$ and $-1-i$; use 20 iterations in each case. Call them, respectively, `root0`, `root1` and `root2`.

<!-- grade_id: cell-73995952add18d91, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print(root0)
print(root1)
print(root2)
```

<!-- grade_id: cell-9fd6f5d355d30747, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4c(_globals)
```

(d) Use NumPy's `vectorize` function to create a version of `newton_sympy`, called `newton_sympy_vec`, which works over arrays. Use it to repeat part (c), returning the roots as a NumPy array `roots` of type complex.

<!-- grade_id: cell-0ee818ce51aa55bc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

roots
```

<!-- grade_id: cell-d952c76e4eb7a907, points: 1, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4d(_globals)
```

(e) Using `vectorize` like that is pretty inefficient. Perform the above task not using `newton_sympy_vec` but instead using `newton_sympy` with `method` set to `lambdify`.

<!-- grade_id: cell-9196113436915df6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

roots
```

<!-- grade_id: cell-eafe5371c23c320a, points: 1, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4e(_globals)
```

(f) Define both `x_values0` and `y_values0` as 1D NumPy arrays of 400 equally-spaced floats between $-2.0$ and $2.0$ inclusive, and use the `meshgrid` function to define `x0` and `y0` as 400-by-400 arrays of, respectively, $x$- and $y$-coordinates specifying a lattice in the plane.

Define `z0` as a lattice of complex numbers of the form $x+i\,y$, with real and imaginary parts taken from your arrays `x0` and `y0`.

Use `newton_sympy` with `method` set to `lambdify` to create an array `roots` consisting of the root of $z^3-1==0$ for each of the elements of the array `z0`, using $100$ iterations in each case. 

Finally, use the `imshow` function from `matplotlib`'s `pyplot` submodule to visualise the imaginary parts of `roots` as they vary over the plane. Use appropriate values for the keyword arguments `origin` and `extent`.

Autograding can't help with that last bit!

<!-- grade_id: cell-68b6cf845f5de10d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-44ad9a945a7b4820, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4f(_globals)
```

## Bonus Question

Write a function that takes in a list of integers, $p(1), p(2), \ldots, p(n)$, where each element is distinct and satisfies $1 \leq p(x) \leq n$. For each $x$ where $1 \leq x \leq n$, the function should find the integer $y$ such that $p(p(y)) = x$ and return the appropriate list. For example, given `p = [5, 2, 1, 3, 4]`, each value of $x$ between 1 and 5 is found as follows:

1. $x = 1 = p(3)$, $p(4) = 3$, $\Rightarrow p(p(4)) = 1$
2. $x = 2 = p(2)$, $p(2) = 2$, $\Rightarrow p(p(2)) = 2$
3. $x = 3 = p(4)$, $p(5) = 4$, $\Rightarrow p(p(5)) = 3$
4. $x = 4 = p(5)$, $p(1) = 5$, $\Rightarrow p(p(1)) = 4$
5. $x = 5 = p(1)$, $p(3) = 1$, $\Rightarrow p(p(3)) = 5$

Therefore, the function should return the list `[4, 2, 5, 1, 3]`.

<!-- grade_id: cell-8f4d566222263187, points: 0, grade: false, solution: true -->
```python
def sequence_explore(p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-043f534519cd6ef3, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = [5, 2, 1, 3, 4]
assert(sequence_explore(test) == [4, 2, 5, 1, 3])
print('Test case passed!!!')
```

<!-- grade_id: cell-0f1099efb70bbc25, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = [4, 3, 5, 1, 2]
assert(sequence_explore(test) == [1, 3, 5, 4, 2])
print('Test case passed!!!')
```

<!-- grade_id: cell-f88bc69d91f086f4, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
```

```python

```